Cryo-EM structure of the bicarbonate receptor GPR30

Author Response:

The following is the authors’ response to the previous reviews

Reviewer #1 (Public review):

Summary:

This study resolves a cryo-EM structure of the GPCR, GPR30, in the presence of bicarbonate, which the author's lab recently identified as the physiological ligand. Understanding the ligand and the mechanism of activation is of fundamental importance to the field of receptor signaling. This solid study provides important insight into the overall structure and suggests a possible bicarbonate binding site.

Strengths:

The overall structure, and proposed mechanism of G-protein coupling are solid. Based on the structure, the authors identify a binding pocket that might accommodate bicarbonate. Although assignment of the binding pocket is speculative, extensive mutagenesis of residues in this pocket identifies several that are important to G-protein signaling. The structure shows some conformational differences with a previous structure of this protein determined in the absence of bicarbonate (PMC11217264). To my knowledge, bicarbonate is the only physiological ligand that has been identified for GPR30, making this study an important contribution to the field. However, the current study provides novel and important circumstantial evidence for the bicarbonate binding site based on mutagenesis and functional assays.

Weaknesses:

Bicarbonate is a challenging ligand for structural and biochemical studies, and because of experimental limitations, this study does not elucidate the exact binding site. Higher resolution structures would be required for structural identification of bicarbonate. The functional assay monitors activation of GPR30, and thus reports on not only bicarbonate binding, but also the integrity of the allosteric network that transduces the binding signal across the membrane. However, biochemical binding assays are challenging because the binding constant is weak, in the mM range.

The authors appropriately acknowledge the limitations of these experimental approaches, and they build a solid circumstantial case for the bicarbonate binding pocket based on extensive mutagenesis and functional analysis. However, the study does fall short of establishing the bicarbonate binding site.

We thank the reviewer for this thoughtful and constructive assessment of our revised manuscript. We are grateful for the recognition of the overall quality of the cryo-EM structure and the proposed mechanism of G-protein coupling, as well as for highlighting the importance of identifying bicarbonate as a physiological ligand for GPR30 and the contribution this work makes to the receptor signaling field. We also appreciate the reviewer’s careful and balanced discussion of the inherent challenges posed by bicarbonate as a low-affinity, small, negatively charged ligand, and we fully agree that, given current experimental limitations, our data provide circumstantial—rather than definitive—evidence for the binding site and that higher-resolution structures would be required for direct visualization. Importantly, we value the reviewer’s acknowledgement that we transparently describe these limitations and that our extensive mutagenesis and functional analyses nonetheless build a solid case for the proposed bicarbonate-binding pocket, which we believe will serve as a useful framework for future biochemical and structural investigation

Reviewer #1 (Recommendations for the authors):

Overall, the authors do a good job responding to the previous review, with updated structures and experimental data. I have two comments on the current version:

(1) When the authors compare their structure to a previously published structure of the same receptor, they say that the previous structure came out while the current manuscript was in revision (line 255). This is not correct. The previous manuscript was published May 14, 2024, and the current manuscript was received by eLife on May 20, 2024. This sentence should be corrected to "During the preparation of this manuscript..."

We corrected the sentence accordingly (line 259).

(2) Line 173: what other structures are the authors referring to? Citations should be included here.

Is Line 193 correct? We added citations (line 190).

Reviewer #2 (Public review):

Summary:

In this manuscript, "Cryo-EM structure of the bicarbonate receptor GPR30," the authors aimed to enrich our understanding of the role of GPR30 in pH homeostasis by combining structural analysis with a receptor function assay. This work is a natural development and extension of their previous work on Nature Communications (PMID: 38413581). In the current body of work, they solved the cryo-EM structure of the human GPR30-G-protein (mini-Gsqi) complex in the presence of bicarbonate ions at 3.15 Å resolution. From the atomic model built based on this map, they observed the overall canonical architecture of class A GPCR and also identified 3 extracellular pockets created by ECLs (Pockets A-C). Based on the polarity, location, size, and charge of each pocket, the authors hypothesized that pocket A is a good candidate for the bicarbonate binding site. To identify the bicarbonate binding site, the authors performed an exhaustive mutant analysis of the hydrophilic residues in Pocket A and analyzed receptor reactivity via calcium assay. In addition, the human GPR30-G-protein complex model also enabled the authors to elucidate the G-protein coupling mechanism of this special class A GPCR, which plays a crucial role in pH homeostasis.

Strengths:

As a continuation of their recent Nature Communications publication, the authors used cryo-EM coupled with mutagenesis and functional studies to elucidate bicarbonate-GPR30 interaction. This work provided atomic-resolution structural observations for the receptor in complex with G-protein, allowing us to explore its mechanism of action, and will further facilitate drug development targeting GPR30. There were 3 extracellular pockets created by ECLs (Pockets A-C). The authors were able to filter out 2 of them and hypothesized that pocket A was a good candidate for the bicarbonate binding site based on the polarity, location, and charge of each pocket. From there, the authors identified the key residues on GPR30 for its interaction with the substrate, bicarbonate. Together with their previous work, they mapped out amino acids that are critical for receptor reactivity.

Weaknesses:

When we see a reduction of a GPCR-mediated downstream signaling, several factors could potentially contribute to this observation: 1) a reduced total expression of this receptor due to the mutation (transcription and translation issue); 2) a reduced surface expression of this receptor due to the mutation (trafficking issue); and 3) a dysfunctional receptor that doesn't signal due to the mutation. In the current revision, based on the gating strategy, the surface expression of the HA-positive WT GPR30-expressing cells is only 10.6% of the total population, while the surface expression levels of the mutants range from 1.89% (P71A) to 64.4% (D111A). Combining this information with the functional readout in Figure 3F and G, as well as their previous work, the authors concluded that mutations at P71, E115, D125, Q138, C207, D210, and H307 would decrease bicarbonate responses. Among those sites,

E115, Q138, and H307 were from their previous Nature Comm paper.

Authors claim P71 and C207 make a structural-stability contribution, as their mutations result in a significant reduction in surface expression: P71A (1.89%) and C207A (2.71%). However, compared to 10.6% of the total population in the WT, (P71A is 17.8% of the WT, and C207A is 25.6% of the WT), this doesn't rule out the possibility that the mutated receptor is also dysfunctional: at 10 mM NaHCO3, RFU of WT is ~500, RFU of P71 and C207 are ~0.

The authors also interpret "The D125ECL1A mutant has lost its activity but is located on the surface" and only mention "D125 is unlikely to be a bicarbonate binding site, and the mutational effect could be explained due to the decreased surface expression". Again, compared to 10.6% of the total population in the WT, D125A (3.94%) is 37.2% of the WT. At 10 mM NaHCO3, the RFU of the WT is ~500, the RFU of D125 is ~0. This doesn't rule out the possibility that the mutated receptor is also dysfunctional. It is not clear why D125A didn't make it to the surface.

Other mutants that the authors didn't mention much in their text: D111A (64.4%, 607.5% of WT surface expression), E121A (50.4%, 475.5% of WT surface expression), R122 (41.0%, 386.8% of WT surface expression), N276A (38.9%, 367.0% of WT surface expression) and E218A (24.6%, 232.1% of WT surface expression) all have similar RFU as WT, although the surface expression is about 2-6 times more. On the other hand, Q215A (3.18%, 30% of WT surface expression) has similar RFU as WT, with only a third of the receptor on the surface.

Altogether, the wide range of surface expression across the different cell lines, combined with the different receptor function readouts, makes the cell functional data only partially support their structural observations.

We sincerely thank the reviewer for their careful reading and thoughtful evaluation of our manuscript on the cryo-EM structure of the bicarbonate receptor GPR30. We greatly appreciate the reviewer’s positive assessment of the overall significance of combining structural determination with extensive mutagenesis and functional assays to advance understanding of bicarbonate–GPR30 interactions and G-protein coupling, as well as their recognition that these atomic-level insights will be valuable for future mechanistic studies and drug-development efforts. We are also grateful for the reviewer’s constructive critique regarding the interpretation of reduced signaling in the context of variable surface expression across mutants, which highlights an important point about disentangling effects of expression/trafficking from intrinsic receptor dysfunction; these comments are highly insightful and will help us strengthen the clarity and rigor of our presentation and conclusions in the revised manuscript.

Reviewer #2 (Recommendations for the authors):

In this revision, the authors have made a significant effort to improve and validate the structural observations, as well as address the comments in the previous submission. They updated the functional assays and evaluated the receptor function by measuring intracellular calcium mobilization, which is a more direct measurement for the downstream signaling of hGPR30-Gq signaling. They also used flow cytometry with an HA-antibody for a more direct measurement of the surface expression of the receptor, replacing their previous assay that normalized to the housekeeping gene Na-K-ATPase.

I appreciate the effort the authors made to address the previous comments made by the reviewers. However, there are still some concerns about the current data.

(1) The authors have addressed my previous comment on untangling the mixture of their previous and new data in the "insights into bicarbonate binding" section. They have made it clear that the importance of E115, Q138, and H307 in the receptor-bicarbonate interaction was shown in their Nature Communications paper.

(2) The authors have addressed my previous comment on adding some content about the physiological concentration of HCO3, or referring more to their previous work about the rationale to select the bicarbonate dose in their functional assay.

(3) The authors have updated Figure 3

(4) The authors have updated Supplemental Figure 1 to show the full gel with molecular weight markers in the supplemental data to demonstrate the sample purity.

(5) The authors have updated the predicted model using AF3

(6) The authors added E218A as suggested before.

Some new suggestions for this R1:

(1) The wide range of surface expression across the different cell lines, combined with the different receptor function readouts, makes the cell functional data only partially support their structural observations.

We acknowledge this limitation. The wide range of surface expression among cell lines, together with differences in assay modalities, may introduce variability that complicates direct quantitative comparisons and therefore only partially supports the structural observations. Future work using more standardized expression systems and matched functional readouts will be important to strengthen the structure–function linkage.

(2) Line 101, "ICL1 and ECL1 contain short α helices", no α helix of ICL1 is shown in Figure 2C

We removed the word “ICL1” (line 98).

(3) For the unsolved region of ECL2, could the author put a dashed line connecting ECL2 with TM4? In the current Figure 2B, it looks like ECL2 connects TM3 and TM5.

According to the suggestion, we corrected Figure 2B.

(4) I appreciate that the authors updated the predicted model with AF3, but they didn't make it clear why they had the comparison between their cryo-EM structure (bicarbonate-activated G-protein-incorporated GPR30) and the predicted AF3 model (inactive GPR30)

We wish to assert the usefulness of experimental structures, not merely predictions. These include structures independent of receptor activation, such as SS bonds.

(5) I appreciate that the authors have addressed my previous comment on adding some content about the physiological concentration of HCO3, but it was still not clear to me why they picked 11 mM in Figure 3G for the bar graph. Also, since a dose-response curve was made in Figure 3F, why not just calculate and report the EC50 of NaHCO3 for each mutant?

Thank you for your comment. Thank you for the comment. We’ve calculated the EC50 of the calcium response and assessed its correlation with receptors’ cell surface expression. We chose 11 mM in Fig .3G since our previous paper in Nature Communications showed the EC50 value of IPs assay was around 11 mM. However, the calcium response was more sensitive and gave a lower value than expected. Therefore, according to your advice, we deleted the bar graph with 11 mM responses, calculated EC50, and drew pictures of the correlation among cell surface expression, EC50, and maximum responses (Figure 3F-I, Supplementary File 1). Moreover, we revised the explanation about this mutagenesis study (lines139-154 and 217-230).

(6) In the previous submission and comments, E218 was in close contact with bicarbonate in the previous Figure 4D (the bicarbonate is deleted in the new structure). I thank the authors for making an E218A mutant and performing the functional assay. As mentioned above, E218A (24.6%, 232.1% of WT surface expression) has a similar functional readout as WT. Doesn't this also indicate that E218A is partially broken, so you will need twice as much as WT to have the same downstream signal?

Thank you for your comment. In our revised manuscript, we described the correlation between cell surface expression and EC50 and found that cell surface expression and the response to bicarbonate are not correlated, which you mentioned in your review comment (Figure 3F-I, Supplementary File 1). There are many possibilities that could explain this: GPR30 localization in specific spots on the plasma membrane might limit the response stoichiometry, GPR30 might also work intracellularly to blunt the increased response because of more GPR30 expression on PM, redundant GPR30 on PM might be broken, or E118A might be less functional and need twice as much as WT. We will examine cell surface expression of GPR30 and its response to bicarbonate in a future study.

I would suggest that the authors in future studies consider using the Tet-on inducible cell lines, such as HEK293 Flp-In Trex. These cell lines will allow the authors to fine-tune the surface expression of their mutants to the same level with different doses of Tetracycline in their stable cell lines.

We appreciate your advice. We’ll introduce Tet-on inducible cell lines for future research.

Reviewer #3 (Public review):

Summary

GPR30 responds to bicarbonate and plays a role in regulating cellular pH and ion homeostasis. However, the molecular basis of bicarbonate recognition by GPR30 remains unresolved. This study reports the cryo-EM structure of GPR30 bound to a chimeric mini-Gq in the presence of bicarbonate, revealing mechanistic insights into its G-protein coupling. Nonetheless, the study does not identify the bicarbonate-binding site within GPR30.

Strengths

The work provides strong structural evidence clarifying how GPR30 engages and couples with Gq.

Weaknesses

Several GPR30 mutants exhibited diminished responses to bicarbonate, but their expression levels were also reduced. As a result, the mechanism by which GPR30 recognizes bicarbonate remains uncertain, leaving this aspect of the study incomplete.

We sincerely thank the reviewer for this thoughtful and balanced assessment of our manuscript, including the clear summary of the central advance and the constructive identification of remaining limitations. We particularly appreciate the recognition that our cryo-EM analysis provides strong structural evidence for how GPR30 engages and couples with Gq, and we agree that pinpointing the bicarbonate-binding site remains a critical open question. In the revised manuscript, we will make this point more explicit, clarify the interpretation of the mutagenesis results in light of reduced receptor expression for some variants, and further strengthen the presentation and discussion of what our current data do—and do not—allow us to conclude regarding bicarbonate recognition by GPR30

Reviewer #3 (Recommendations for the authors):

The authors have removed the bicarbonate assignment from their model and have addressed all of my concerns. In this study, or in future work, it would be advisable for the authors to explore the use of bicarbonate mimetics with higher binding affinity to facilitate more definitive structural characterization.

Thank you for this constructive suggestion. We agree that exploring bicarbonate mimetics with higher binding affinity would be an important next step to enable more definitive structural characterization of GPR30 and to strengthen mechanistic conclusions. In future work, we plan to pursue the identification and/or design of such mimetics, guided by the architecture and mutational landscape of the extracellular pocket described here, and to combine these ligands with optimized cryo-EM sample preparation and complementary functional assays to better stabilize and visualize the bound state.

Changes from the earlier version to the revised version

Abstract

G-protein-coupled receptor 30 (GPR30) is a bicarbonate receptor that plays a vital role in cellular responses to extracellular pH and ion homeostasis. Despite its significance, the mechanisms by which GPR30 interacts with bicarbonate ions remain elusive. There is no consensus on a drug that targets GPR30, and difficulties in pharmacological analyses have limited biological and drug discovery research on GPR30. Here, we present the cryo-electron microscopy structure of human GPR30 in the presence of bicarbonate ions, at 3.15 Å resolution. Our structure reveals unique extracellular pockets and critical residues for bicarbonate binding and activation. Functional assays demonstrate that mutations in these residues impair bicarbonate-induced GPR30 activation, underscoring their importance in receptor function. This study also provides insights into G-protein coupling, highlighting the structural divergence between GPR30 and other GPCRs. Our findings not only advance the understanding of the role of GPR30 in pH homeostasis but also pave the way for the development of high-affinity drugs targeting GPR30 for therapeutic interventions in diseases associated with acid-base imbalance.

Introduction

Earlier version: The extracellular environment regulates cellular responses, as exemplified by pH and ion homeostasis. The acid-base balance, largely based on the bicarbonate buffer system in vivo, serves as a vital mechanism that maintains the optimal pH for cellular function.1 Changes in the extracellular pH and ion homeostasis are monitored by membrane channels and receptors, such as G protein-coupled receptors (GPCRs). However, it remains unknown whether and how the acid–base balance-related ions, protons and bicarbonate ions, bind to the receptors and cause the conformational changes that lead to intracellular signal transduction. These proton-sensing GPCRs are thought to be involved in pH homeostasis, particularly in the acidic tumor microenvironment, at inflamed sites, and during ischemia-reperfusion injury2. In recent years, the determination of several proton receptor structures has begun to elucidate the molecular mechanisms underlying proton sensing3–7.

Revised version: The extracellular environment regulates cellular responses, as exemplified by pH and ion homeostasis. The acid-base balance, largely based on the bicarbonate buffer system in vivo, serves as a vital mechanism that maintains the optimal pH for cellular function1. Changes in the extracellular pH and ion homeostasis are monitored by membrane channels and receptors, such as G protein-coupled receptors (GPCRs). However, it remains unknown whether and how the acid–base balance-related ions, protons and bicarbonate ions, bind to the receptors and cause the conformational changes that lead to intracellular signal transduction. These proton-sensing GPCRs are thought to be involved in pH homeostasis, particularly in the acidic tumor microenvironment, at inflamed sites, and during ischemia-reperfusion injury2. In recent years, the determination of several proton receptor structures has begun to elucidate the molecular mechanisms underlying proton sensing3–7.

We recently reported that a physiological concentration of bicarbonate ions, the counterpart of protons in the bicarbonate buffer system, activates G protein-coupled receptor 30 (GPR30), which leads to Gq-coupled calcium responses8. Our study also demonstrated that GPR30 in brain mural cells regulates blood flow and ischemia–reperfusion injury. GPR30 was identified as a G protein-coupled estrogen receptor that mediates the rapid non-genomic action of estradiol (E2)9. However, despite numerous reports on the pleiotropic functions of GPR30 in vivo10,11, controversy remains regarding the responses of GPR30 to E2 in vitro12, ex vivo13, and in vivo14. The broad expression of GPR3015,16, including blood vessels17,18, stomach, and lung, has also raised the possibility of its non-estrogenic functions. We recently demonstrated that three amino acids, E1152.61, Q1383.33, and H3077.36, are essential for the bicarbonate-induced activation of GPR30, according to the public homology model (https://gpcrdb.org/)19. However, we did not examine whether bicarbonate ions interact with GPR30 and cause conformational changes. Moreover, to date, there is no consensus on a high-affinity drug that targets GPR30, and difficulties in pharmacological analyses have limited biological and drug discovery research on GPR30. To elucidate the bicarbonate–GPR30 interaction and to establish the first structural identification of an acid-base balance-related GPCR, we report the cryo-electron microscopy (cryo-EM) structure of human GPR30 in the presence of bicarbonate ions.

Results

Overall structure

For the structural study, we used the full-length human GPR30 sequence. To efficiently purify the stable GPCR-G-protein complex, the receptor and mini-Gsqi were incorporated in a ‘Fusion-G system’ by combining two complex stabilization techniques20,21 (Figure 1—figure supplement 1A, B). The modified receptor and G-protein were co-expressed in Sf9 cells and purified by Flag affinity chromatography in the presence of 200 mM NaHCO3. After an incubation with nanobody 35 (Nb35) and single chain scFv16, which binds to mini-Gsqi, the complex was purified by size exclusion chromatography (Figure 1—figure supplement 1C, D). The structure of the purified complex was determined by single-particle cryo-EM analysis, with an overall resolution of 3.15 Å (Table 1, Figure 1—figure supplement 2, and “Methods”). As the extracellular portion of the receptor was poorly resolved, we performed receptor focused refinement, yielding a density map with a nominal resolution of 3.18 Å, which was combined with the refined map focused on G-protein. The resulting composite map allowed us to precisely build the atomic model of the components, including the receptor (residues 51 to 196, 207 to 288, and 299 to 340, G-proteins, and scFv16) (Figure 1A, B). Nb35 was not visible in the cryo-EM map, probably because of the effect of the mini-Gs modification, although the interactive residues with Nb35 were not changed.

Overall structure of the GPR30-miniGsqiβ1γ2-scFv16 complex.

A Unsharpened cryo-EM density map of the GPR30-miniGsqiβ1γ2-scFv16-Nb35 complex, with the components individually colored. B The refined structure of the complex is shown as a ribbon representation.

Earlier version: The receptor consists of the canonical 7 transmembrane helices (TM) connected by three intracellular loops (ICL1–3) and three extracellular loops (ECL1–3), and the amphipathic helix 8 at the C-terminus (H8) (Figure 2A). Most of the TMs are kinked, as in typical GPCRs22. It should be noted that TM1 is also kinked at P711.44 (superscripts indicate Ballesteros-Weinstein numbers23), whereas TM1 adopts a straight α-helix in most class A GPCRs. ICL1 and ECL1 contain short α helices (Figure 2B, C). The short ICL3 was completely visible in the cryo-EM map (Figure 2C and Figure 1—figure supplement 2), while residues 197 to 206 in ECL2 and 289 to 298 in ECL3 were disordered (Figure 2B). ECL2 is attached to TM3 by the disulfide bond between C1303.25 and C207ECL2 (Figure 2B, Figure 2—figure supplement 1A-F), which is highly conserved in class A GPCRs22,24. The cryo-EM structure did not superimpose well on the AlphaFold-predicted structure (Q99527 in the AlphaFold database)25, with a root mean square deviation (R.M.S.D.) of Cα atoms of 3.10 Å (Figure 2—figure supplement 2A). The conserved D3.49R3.50Y3.51 motif in the predicted structure represents an inactive state (Figure 2—figure supplement 1B)22. Moreover, ECLs 2 and 3 are rich in cysteine residues, which form incorrect disulfide bonds in the predicted structure (Figure 2— figure supplement 2C). These comparisons support the usefulness of experimental structural determination.

Revised version: The receptor consists of the canonical 7 transmembrane helices (TM) connected by three intracellular loops (ICL1–3) and three extracellular loops (ECL1–3), and the amphipathic helix 8 at the C-terminus (H8) (Figure 2A). Most of the TMs are kinked, as in typical GPCRs22. It should be noted that TM1 is also kinked at P711.44 (superscripts indicate Ballesteros-Weinstein numbers23), whereas TM1 adopts a straight α-helix in most class A GPCRs. ECL1 contains a short α helix (Figure 2B). The short ICL3 was completely visible in the cryo-EM map (Figure 2C and Figure 1—figure supplement 2), while residues 197 to 206 in ECL2 and 289 to 298 in ECL3 were disordered (Figure 2B). ECL2 is attached to TM3 by the disulfide bond between C1303.25 and C207ECL2 (Figure 2B, Figure 2—figure supplement 1A-F), which is highly conserved in class A GPCRs22,24. The cryo-EM structure did not superimpose well on the AlphaFold-predicted structure (Q99527 in the AlphaFold database)25, with a root mean square deviation (R.M.S.D.) of Cα atoms of 3.10 Å (Figure 2—figure supplement 2A). The conserved D3.49R3.50Y3.51 motif in the predicted structure represents an inactive state (Figure 2—figure supplement 1B)22. Moreover, ECLs 2 and 3 are rich in cysteine residues, which form incorrect disulfide bonds in the predicted structure (Figure 2—figure supplement 2C). These comparisons support the usefulness of experimental structural determination.

Receptor structure.

A–C Overall structure of the receptor, viewed from the membrane plane (A), intracellular side (B), and extracellular side (C).

Extracellular pockets

Earlier version: The interaction network between ECL1–3 covers the extracellular side of the receptor. As described above, ECL2 is anchored to TM3 via the disulfide bond, and ECL3 extends between ECL1 and ECL2, with R299ECL3 located within the receptor cavity (Figure 3A, B). R299ECL3 forms an electrostatic interaction with D210ECL2. These interactions among the ECLs create three extracellular pockets (pockets A–C) (Figure 3A). Pocket B consists of ECL1, TM1, and TM7 (Figure 3C), while pocket C is formed by TM6, TM7, and ECL1 (Figure 3D). These two pockets are small, superficial, and hydrophobic and appear unsuitable for bicarbonate binding. Pocket A consists of ECL1–3 and TM2–7 and is the largest among the three pockets. It is connected to the inside of the receptor, and deep enough to reach W2726.48 (Figure 3B). Pocket A contains numerous hydrophilic residues favorable for bicarbonate binding (Figure 3E), suggesting that pocket A is a good candidate for the bicarbonate binding site.

Revised version: The interaction network between ECL1–3 covers the extracellular side of the receptor. As described above, ECL2 is anchored to TM3 via the disulfide bond, and ECL3 extends between ECL1 and ECL2, with R299ECL3 located within the receptor cavity (Figure 3A, B). R299ECL3 forms an electrostatic interaction with D210ECL2. These interactions among the ECLs create three extracellular pockets (pockets A–C) (Figure 3A). Pocket B consists of ECL1, TM1, and TM7 (Figure 3C), while pocket C is formed by TM4, TM5, and ECL1 (Figure 3D). These two pockets are small, superficial, and hydrophobic and appear unsuitable for bicarbonate binding. Pocket A consists of ECL1–3 and TM2–7 and is the largest among the three pockets. It is connected to the inside of the receptor, and deep enough to reach W2726.48 (Figure 3B). Pocket A contains numerous hydrophilic residues favorable for bicarbonate binding (Figure 3E), suggesting that pocket A is a good candidate for the bicarbonate binding site.

Architecture of the extracellular pocket.

Earlier version: A Molecular surface of the extracellular side. B Cross section of the pocket. C–E Residues facing pocket B (C), pocket C (D), and pocket A (E). In panel (E), only residues with reduced bicarbonate responses are highlighted. F, G Calcium assay using stable HEK293 cell lines expressing the N-terminal HA-tagged GPR30 wild-type (WT) and mutants. The mutants D1112.56A, N2766.52A, and Q296ECL3A; E121ECL1A, R122ECL1A, S1343.29A, Q2155.39A, and E2185.42A; D210ECL2A and Q1383.33A; and D125ECL1A, C207ECL2A, and P711.44A are highlighted in red, blue, green, and purple, respectively. The cells were stimulated by the indicated concentrations of NaHCO3 at the timepoint of t = 20 sec. The Y axis indicates the difference between the maximum and minimum fluorescent values during 15 to 60 sec. Statistical analysis: each sample was compared with WT cells using Dunnett’s T3 multiple comparisons test after Brown-Forsythe and Welch ANOVA tests. **** p < 0.0001 compared to WT cells, no significant difference is not shown. Data are presented as mean values ± SEM (F) and each replicate with the mean ± SEM (G).

Revised version: A Molecular surface of the extracellular side. B Cross section of the pocket. C–E Residues facing pocket B (C), pocket C (D), and pocket A (E). In panel (E), only residues with reduced bicarbonate responses are highlighted. F Calcium assay using Stable HEK293 cell lines expressing the N-terminal HA-tagged wild-type (WT) and mutant GPR30. The mutants D1112.56A, N2766.52A, Q296ECL3A; E121ECL1A, R122ECL1A, S1343.29A, Q2155.39A, E2185.42A; D210ECL2A, Q1383.33A; and D125ECL1A, C207ECL2A, P711.44A are highlighted in red, blue, green, and purple, respectively. The cells were stimulated by the indicated concentrations of NaHCO3 and 50 µM ATP at the timepoint of t = 20 sec. The Y axis indicates the difference between the maximum and minimum fluorescent values during 15 to 60 sec, normalized by those with ATP stimulation. Nonlinear regression (four-parameter) was used for curve-fitting. G–H X-Y plot of parameters calculated by FACS analysis and Ca assay. D1112.56A; D210ECL2A, Q1383.33A; D125ECL1A, C207ECL2A, P711.44A; and mock are colored in red, green, purple, and grey, respectively(G). cell surface (x-axis) and total (y-axis) expression of GPR30 (H). X-Y plot of cell surface (x-axis) and normalized maximum responses (y-axis) (I). X-Y plot of cell surface (x-axis) and EC50 (y-axis). EC50 was calculated using the data shown in panel (F).

However, no obvious density corresponding to bicarbonate is observed in pocket A, along with pockets B and C. In cryo-EM, negative charges are generally difficult to visualize, and the local resolution of the extracellular region in this structure is moderate (Figure 1—figure supplement 2). Thus, even if bicarbonate ions are present, it is possible that they cannot be visualized.

Earlier version: Our previous studies have shown that the mutations of E1152.61, Q1383.33, and H3077.36 abolish the bicarbonate-dependent calcium response of GPR30, suggesting their contribution to bicarbonate binding. To predict the bicarbonate binding site, we performed an exhaustive mutant analysis of the hydrophilic residues in pocket A, which was not done in the previous study8. We performed calcium assays using cell lines stably expressing N-terminally HA-tagged GPR30 and its mutants (Figure 3F, G), whose surface expression levels were analyzed by fluorescence-activated cell sorting (FACS) using an HA-antibody (Figure 3—figure supplement 1-3). Moreover, we performed TGFα shedding assays26 using HEK cells transiently expressing the receptors (Figure 3—figure supplement 4A-F). The HA-tagged GPR30 wild-type showed a bicarbonate-dependent response. By contrast, the C207ECL2A mutation that disrupts the disulfide bond between ECL2 and TM3, and the P711.44 mutation that disrupts the characteristic kink in TM1, abolished the bicarbonate-dependent response. The expression levels of these mutants on the cell surface are significantly reduced (Figure 3—figure supplement 1, 4F), indicating that their reduced calcium responses are due to indirect effects such as structural instability, rather than the direct binding of bicarbonate ions. These findings underscore the critical role of the structural integrity of the receptor and provide compelling evidence that the bicarbonate response is mediated through GPR30.

Revised version: Our previous studies have shown that the mutations of E1152.61, Q1383.33, and H3077.36 abolish the bicarbonate-dependent calcium response of GPR30, suggesting their contribution to bicarbonate binding. To predict the bicarbonate binding site, we performed an exhaustive mutant analysis of the hydrophilic residues in pocket A, which was not done in the previous study8. We performed calcium assays using cell lines stably expressing N-terminally HA-tagged GPR30 and its mutants (Figure 3F), whose surface expression levels were analyzed by fluorescence-activated cell sorting (FACS) using an HA-antibody (Figure 3—figure supplement 1-3). We also performed TGFα shedding assays26 using HEK cells transiently expressing the receptors (Figure 3—figure supplement 4A-F).

Added in the revised version: The D125 ECL1A, C207ECL2A, and P711.44 A mutations abolished the bicarbonate-dependent calcium responses. The Q1383.33A and D210 ECL2A mutations reduced reactivity (Figure 3F). Cell surface expression levels of wild type and mutant GPR30 ranged widely, though the non- and little-responsive mutants, C207ECL2A, P711.44A, D125 ECL1A, Q1383.33A, and D210 ECL2A showed minimum cell surface expression compared with the others (Figure 3G). We then assessed the correlation between cell surface expression and calcium responses: maximum response (Figure 3H) and EC50 (Figure 3I), to examine if the mutants’ reactivity to bicarbonate was attributed to their expression levels on the plasma membrane. Cell surface expression and responsiveness to bicarbonate were not correlated. Q1383.33A and D210 ECL2A mutants showed one-third and half maximum responses of GPR30 wild-type, respectively (Figure 3H). The Q1383.33A mutant showed a higher EC50 value than wild type GPR30, which suggests the involvement in bicarbonate-binding (Figure 3I). D210 ECL2A mutation impairs downstream signaling while keeping its affinity to bicarbonate (Figure 3I). In contrast, the D1112.56A mutation enhanced receptor affinity, presenting the half value of wild type EC50 (Figure 3.—figure Supplement 4).

Removed in the revised version: Among the hydrophilic residue mutants, the D125ECL1A mutant abolished the bicarbonate-induced responses (Figure 3F, G, and Figure 3—figure supplement 4A, B). Q1383.33A and D210ECL2A reduced the responses and showed significant reductions in their cell surface expression levels (Figure 3F, G, and Figure 3—figure supplement 1, 4A, B, F). However, caution is required in interpreting the results, since the cell surface expression level of the Q2155.39A mutant was reduced to the same extent as Q1383.33A and D210ECL2A, but its bicarbonate-induced calcium response was as high as that of WT (Figure 3—figure supplement 1).

Earlier version: Combining these data with previous studies, we mapped the residues that cause decreased bicarbonate responses due to mutations on the current structure (P711.44, E1152.60, D125ECL1, Q1383.33, C207ECL2, D210ECL2, and H3077.36) (Figure 3E). Except for the structure-contributing residues (C207ECL2 and P711.44), the essential residues face pocket A, further supporting its role as the bicarbonate binding site. The D210ECL2A mutant retained a moderate bicarbonate-induced calcium response (Figure 3F, G) and would contribute to the structural stability of the pocket through its interaction with R299ECL3 (Figure 3A). The D125ECL1A mutant has lost its activity but is located on the surface (Figure 3E-G). Thus, D125ECL1 is unlikely to be a bicarbonate binding site, and the mutational effect could be explained due to the decreased surface expression (Figure 3—figure supplement 1, 4F). Given that most bicarbonate binding modes in other structures require recognition by positively charged residues, bicarbonate may bind to H3077.36, while E1152.61, which is located near H3077.36, may play an essential role by coordinating cations for bicarbonate binding, as observed in other structures. These residues are evolutionarily conserved from fish to human, further supporting their importance (Figure 3—figure supplement 5).

Revised version: Combining these data with previous studies, we mapped the residues that cause decreased bicarbonate responses due to mutations on the current structure (P711.44, E1152.60, D125ECL1, Q1383.33, C207ECL2, D210ECL2, and H3077.36) (Figure 3E). Except for the structure-contributing residues (C207ECL2 and P711.44), the essential residues face pocket A, further supporting its role as the bicarbonate binding site. The D210ECL2A mutant retained a moderate bicarbonate-induced calcium response (Figure 3F, G) and would contribute to the structural stability of the pocket through its interaction with R299ECL3 (Figure 3A). The D125ECL1A mutant has lost its activity but is located on the surface (Figure 3E-G) and is unlikely to be a bicarbonate binding site. Given that most bicarbonate binding modes in other structures require recognition by positively charged residues, bicarbonate may bind to H3077.36, while E1152.61, which is located near H3077.36, may play an essential role by coordinating cations for bicarbonate binding, as observed in other structures. These residues are evolutionarily conserved from fish to human, further supporting their importance (Figure 3—figure supplement 5).

G-protein coupling

The C-terminal helix of Gαq (α5-helix) deeply enters the intracellular cavity of GPR30, resulting in the formation of an active signaling complex. On the C-terminal side of the α5-helix, the backbone carbonyl of Y243G.H5.23 (superscript indicates the common Gα numbering [CGN] system27) hydrogen bonds with R1553.50 in the conserved D3.49R3.50Y3.51 motif22 (Figure 4A). In addition, N244G.H5.24 forms a hydrogen bond with the peptide backbone of Y3247.53. ICL3 faces the middle part of the α5-helix, with three characteristic arginine residues (R2486.24, R2516.27, and R2546.30) surrounding D233G.H5.13 (Figure 4B). R2486.24 and R2546.30 form hydrogen bonds with D233G.H.13 and Q237G.H5.17, respectively. At ICL2, M16334.51 fits into a hydrophobic pocket composed of the α5-helix, the αN-β2 loop, and the β2-β3 loop of the Gα subunit, as in other GPCR-Gq complexes28–32 (Figure 4C–H). Above it, L1593.54 and A16234.50 form extensive hydrophobic interactions with L235G.H5.15, L236G.H5.16, and L240G.H5.20 in the α5-helix. These tight interactions with the α5-helix would enable the coupling of GPR30 with Gq.

G-protein coupling.

A Hydrogen-bonding interactions between the C-terminal α5-helix residues and the receptor. B Electrostatic and hydrogen-binding interactions between ICL3 and the α5-helix. C–H Interface between ICL2 and Gq, with residues involved in hydrophobic interactions represented by CPK models. G-protein-bound GPCRs used in the comparison are as follows: H1R-Gq (PDB 7DFL, gray), B2R-Gq (PDB 7F6I, yellow-green), MRGPRX2-Gq (PDB 7S8L, green), 5-HT2A-Gq (PDB 6WHA, gray), and GPR103-Gq (PDB 8ZH8, red). I Comparison of the angles and positions of α5h and αN relative to the receptor. J Superimposition of the Gα subunits.

Earlier version: The position of the α5-helix aligns with those in other GPCR-Gq complexes, but the position of the αN-helix is rotated away from the receptor (Figure 4I). The superimposition of the Gα subunits indicates that the rotation is not due to the movement of the entire G protein, but rather a 33° bend in the αN-helix (Figure 4J). This is attributed to variations in the interactions of the αN-helix with its receptors. The αN-helix interacts with the intracellular ends of ICL2 and TM4 in receptor-specific modes. In most Gq-coupled structures, ICL2 forms an α-helix and R34.55 interacts with R38G.hns1.03 at the end of the αN-helix. A typical example is the head-to-toe interaction between R18534.55 and R32G.hns1.03 in the serotonin 5-HT2A receptor31 (Figure 4E). However, in GPR30, ICL2 does not adopt an α-helix and R34.55 is replaced by L16734.55, which displaces R32G.hns1.03 downward through their interaction. Moreover, T17034.58 and R16934.57 form direct hydrogen bonds with R31G.hns1.02 and Y243G.H5.23, respectively, which are not observed in other GPCR-Gq complexes. We hypothesize that these interactions may underpin the distinct αN rotation: A similar interaction between R38G.hns1.03 and L13634.54 is also observed in histamine H1 receptor (H1R)28 (Figure 4D), in which translational rather than rotational motion of αN is observed (Figure 4J).

Revised version: The position of the α5-helix aligns with those in other GPCR-Gq complexes, but the position of the αN-helix is rotated away from the receptor (Figure 4I). The superimposition of the Gα subunits indicates that the rotation is not due to the movement of the entire G protein, but rather a 33° bend in the αN-helix (Figure 4J). This is attributed to variations in the interactions of the αN-helix with its receptors. The αN-helix interacts with the intracellular ends of ICL2 and TM4 in receptor-specific modes. In most Gq-coupled structures28–32, ICL2 forms an α-helix and R34.55 interacts with R38G.hns1.03 at the end of the αN-helix. A typical example is the head-to-toe interaction between R18534.55 and R32G.hns1.03 in the serotonin 5-HT2A receptor31 (Figure 4E). However, in GPR30, ICL2 does not adopt an α-helix and R34.55 is replaced by L16734.55, which displaces R32G.hns1.03 downward through their interaction. Moreover, T17034.58 and R16934.57 form direct hydrogen bonds with R31G.hns1.02 and Y243G.H5.23, respectively, which are not observed in other GPCR-Gq complexes. We hypothesize that these interactions may underpin the distinct αN rotation: A similar interaction between R38G.hns1.03 and L13634.54 is also observed in histamine H1 receptor (H1R)28 (Figure 4D), in which translational rather than rotational motion of αN is observed (Figure 4J).

Discussion

Earlier version: Our study has elucidated the distinctive structure of GPR30, which is unique among the class A GPCRs. The GPR30 structure is characterized by the kink of TM1 and the presence of the large extracellular pocket covered by ECLs. This pocket is rich in hydrophilic residues and could be a focus for the development of high-affinity drugs targeting GPR30. Our comprehensive mutagenesis study of the binding pocket complemented our previous study and mapped the essential residues for the bicarbonate-induced calcium response. Mutations in the S-S bond between ECL2 and TM3 and the bend in TM1 support their importance in the stability and the shape of the GPR30 structure and pocket. Furthermore, we were able to identify residues in pocket A that are important for the bicarbonate response, although there is room for debate regarding whether they directly contribute to binding or their mutations caused significant reductions in membrane expression. Considering the binding sites of negatively charged bicarbonate in other protein structures, it is reasonable to regard the region around H3077.36 as a bicarbonate binding site candidate. To clarify the binding mode of bicarbonate, a higher resolution structural analysis and complementary functional assessments are necessary.

Revised version: Our study has elucidated the distinctive structure of GPR30, which is unique among the class A GPCRs. The GPR30 structure is characterized by the kink of TM1 and the presence of the large extracellular pocket covered by ECLs. This pocket is rich in hydrophilic residues and could be a focus for the development of high-affinity drugs targeting GPR30. Our comprehensive mutagenesis study of the binding pocket complemented our previous study and mapped the essential residues for the bicarbonate-induced calcium response. Mutations in the S-S bond between ECL2 and TM3 and the bend in TM1 support their importance in the stability and the shape of the GPR30 structure and pocket. Furthermore, we were able to identify residues in pocket A that are important for the bicarbonate response, although there is room for debate regarding whether they directly contribute to binding or compromise the structure and thereby downstream signaling. Considering the binding sites of negatively charged bicarbonate in other protein structures, it is reasonable to regard the region around H3077.36 as a bicarbonate binding site candidate. To clarify the binding mode of bicarbonate, a higher resolution structural analysis and complementary functional assessments are necessary.

Added in the revised version: The C207ECL2A and P711.44 mutation disrupts the disulfide bond between ECL2 and TM3 and the characteristic kink in TM1, respectively. Loss of bicarbonate-dependent activation by these mutations underscores the critical role of the structural integrity of the receptor and provides compelling evidence that the bicarbonate response is mediated through GPR30. Higher EC50 value of Q1383.33A mutant than that of wild type suggests that the residue might be involved in bicarbonate-binding. D210 ECL2A mutation impairs downstream signaling while keeping its affinity to bicarbonate. Overall, cell surface expression and response to bicarbonate were not correlated as long as only a small portion (∼10%) was expressed on the cell surface. Many possibilities could explain this: GPR30 localization in specific spots on the plasma membrane might limit the response stoichiometry, GPR30 might also work intracellularly to blunt the increased response because of more GPR30 expression on the plasma membrane, redundant GPR30 on the plasma membrane might be broken, or the mutants might be less functional than wild type receptor and need higher expression.

GPR30 exhibits low homology to other GPCRs, with even closely related receptors sharing less than 30% sequence identity. A BLAST analysis revealed that GPR30 shares the highest sequence identity of 28% with the type 2 angiotensin II receptor (AT2R) (Figure 5—figure supplement 1). Consequently, we compared the current structure of GPR30 with the ligand-bound AT2R33. The superimposition of GPR30 and AT2R, with a Cα R.M.S.D. of 1.73 Å, suggests significant structural divergence between the two (Figure 5A). Notably, there is no structural commonality between GPR30 and AT2R in terms of the orientations of their respective TMs. Specifically, the extracellular half of TM1 in GPR30 extends outward, and is kinked at P711.44 in the center of TM1 (Figure 5B). Near P711.44, Y651.38 and F701.43 are oriented toward the interior of the receptor, forming hydrophobic interactions with L1132.58 and L3317.40 to stabilize the kink. In the examination of the loops, ECL2 of AT2R features a short β-sheet commonly found in peptide-activated class A GPCRs34–36, whereas that of GPR30 adopts an elongated conformation that covers the extracellular pocket essential for the bicarbonate response (Figure 5A, B). These comparisons underscore the unique structural characteristics of GPR30.

Structural comparison with related GPCRs.

A Structural comparison of GPR30 and AT2R (PDB 5UNF). B Interactions around P711.44 in TM1 of GPR30. C–F Conformational changes of TM1 upon agonist binding in CB1 (C), A2AR (D), and ETB (E). The agonist-bound states are colored with the respective colors, while the inactive states are colored gray. The PDB codes used in this figure are CB1-active (PDB 5XRA), CB1-inactive (PDB 5TGZ), A2AR-active (PDB 6GDG), A2AR-inactive (PDB 3EML), ETB-active (PDB 8IY5), and ETB-inactive (PDB 5X93).

In most GPCRs, TM1 adopts a straight α-helix and is positioned away from the orthosteric pocket formed by TM2–7. However, TM1 undergoes an inward movement due to allosteric conformational changes during agonist binding. The cannabinoid receptor CB1 and adenosine A2A receptor (A2AR) exemplify this phenomenon, with TM1 kinked at G1271.43 and G231.49 folding inward upon ligand binding37–40 (Figure 5C, D). By contrast, TM1 moves 4 Å outward in the endothelin ETB receptor41,42 (Figure 5E). In GPR30, TM1 is kinked at P711.44 (Figure 5B), indicating that this kink exists regardless of the receptor activation state. Indeed, the P711.44A mutation, which eliminates the kink in TM1, completely abolished the bicarbonate response (Figure 3F, G). Moreover, P711.44 is evolutionarily conserved (Figure 3—figure supplement 5). These findings suggest the critical role of the TM1 kink in both the putative conformational changes associated with ligand binding and the functional integrity of GPR30.

Earlier version: During the revision of this manuscript, the structures of apo-GPR30-Gq (PDB 8XOG) and the exogenous ligand Lys05-bound GPR30-Gq (PDB 8XOF) were reported43. We compared our structure of GPR30 in the presence of bicarbonate with these structures. In the extracellular region, the position of TM5 in GPR30 in the presence of bicarbonate is similar to that in apo-GPR30. In contrast, the position of TM6 is shifted outward relative to that of apo-GPR30, resembling the conformation observed in Lys05-bound GPR30 (Figure 6A, B). Additionally, the position of ECL1 is also shifted outward compared to that of apo-GPR30 (Figure 6B). In the GPR30 structure in the presence of bicarbonate, ECL2 was modeled, suggesting differences in structural flexibility. These findings indicate that the structure of GPR30 in the presence of bicarbonate is different from both the apo structure and the Lys05-bound structure, demonstrating that the structure and the flexibility of the extracellular domain of GPR30 change depending on the type of ligand. Furthermore, focusing on the interaction with Gq, the αN helix of Gq is not rotated in the structure bound to Lys05, in contrast to the characteristic bending of the αN helix in our structure (Figure 6C, D). Although it is necessary to consider variations in experimental conditions, such as salt concentration, the differences in the Gq binding modes suggest that the downstream signals may change in a ligand-dependent manner.

Revised version: During the preparation of the manuscript, the structures of apo-GPR30-Gq (PDB 8XOG) and the exogenous ligand Lys05-bound GPR30-Gq (PDB 8XOF) were reported43. We compared our structure of GPR30 in the presence of bicarbonate with these structures. In the extracellular region, the position of TM5 in GPR30 in the presence of bicarbonate is similar to that in apo-GPR30. In contrast, the position of TM6 is shifted outward relative to that of apo-GPR30, resembling the conformation observed in Lys05-bound GPR30 (Figure 6A, B). Additionally, the position of ECL1 is also shifted outward compared to that of apo-GPR30 (Figure 6B). In the GPR30 structure in the presence of bicarbonate, ECL2 was modeled, suggesting differences in structural flexibility. These findings indicate that the structure of GPR30 in the presence of bicarbonate is different from both the apo structure and the Lys05-bound structure, demonstrating that the structure and the flexibility of the extracellular domain of GPR30 change depending on the type of ligand. Furthermore, focusing on the interaction with Gq, the αN helix of Gq is not rotated in the structure bound to Lys05, in contrast to the characteristic bending of the αN helix in our structure (Figure 6C, D). Although it is necessary to consider variations in experimental conditions, such as salt concentration, the differences in the Gq binding modes suggest that the downstream signals may change in a ligand-dependent manner.

Structural comparison with GPR30 bound to Lys05.

A, B Superimposition of bicarbonate-GPR30, Lys05-binding GPR30 (PDB 8XOF) (A), and apo-GPR30 (PDB 8XOG) (B) structures, viewed from the extracellular side. C Comparison of the angles and positions of α5h and αN relative to the receptor. D Superimposition of the Gα subunits.

In our work, we used relatively high (mM) concentrations of bicarbonate to activate GPR30. In parallel with its low potency, the physiological concentrations of bicarbonate are 22-26 mM in the extracellular fluid, including interstitial fluid and blood, and 10-12 mM in the cytoplasm. Therefore, GPR30 is activated by the physiological concentrations of bicarbonate in tissues. The bicarbonate concentrations are altered in various physiological and pathological conditions – metabolic acidosis in chronic kidney disease causes a drop to 2-3 mM, and metabolic alkalosis induced by severe emesis increases bicarbonate concentrations to over 30 mM. Thus, our work clearly shows that GPR30 is activated by physiological concentrations of bicarbonate, present both intracellularly and extracellularly, and that GPR30 can be deactivated or reactivated under various pathophysiological conditions. Thus, we have identified GPR30 as a bicarbonate receptor.

Methods

Expression and purification of scFv16 and Nb35

Earlier version: The gene encoding scFv16 was synthesized (GeneArt) and subcloned into a modified pFastBac vector44, with the resulting construct encoding the GP67 secretion signal sequence at the N terminus, and a His8 tag followed by a TEV cleavage site at the C terminus. The His8-tagged scFv16 was expressed and secreted by Sf9 insect cells, as previously reported45. The Sf9 cells were removed by centrifugation at 5,000g for 10 min, and the secreta-containing supernatant was combined with 5 mM CaCl2, 1 mM NiCl2, 20 mM HEPES (pH 8.0), and 150 mM NaCl. The supernatant was mixed with Ni Superflow resin (GE Healthcare Life Sciences) and stirred for 1 h at 4 ℃. The collected resin was washed with buffer containing 20 mM HEPES (pH 8.0), 500 mM NaCl and 20 mM imidazole, and further washed with 10 column volumes of buffer containing 20 mM HEPES (pH 8.0), 500 mM NaCl and 20 mM imidazole. Next, the protein was eluted with 20 mM Tris (pH 8.0), 500 mM NaCl and 400 mM imidazole. The eluted fraction was concentrated and loaded onto a Superdex200 10/300 Increase size-exclusion column, equilibrated in buffer containing 20CmM Tris (pHC8.0) and 150CmM NaCl. Peak fractions were pooled, concentrated to 5Cmg/ml using a centrifugal filter device (Millipore 10CkDaCMW cutoff), and frozen in liquid nitrogen.

Revised version: The gene encoding scFv16 was synthesized (GeneArt) and subcloned into a modified pFastBac vector44, with the resulting construct encoding the GP67 secretion signal sequence at the N terminus, and a His8 tag followed by a TEV cleavage site at the C terminus. The His8-tagged scFv16 was expressed and secreted by Sf9 insect cells, as previously reported45. The Sf9 cells were removed by centrifugation at 5,000g for 10 min, and the secreta-containing supernatant was combined with 5 mM CaCl2, 1 mM NiCl2, 20 mM HEPES (pH 8.0), and 150 mM NaCl. The supernatant was mixed with Ni Superflow resin (GE Healthcare Life Sciences) and stirred for 1 h at 4 . The collected resin was washed with buffer containing 20 mM HEPES (pH 8.0), 500 mM NaCl and 20 mM imidazole, and further washed with 10 column volumes of buffer containing 20 mM HEPES (pH 8.0), 500 mM NaCl and 20 mM imidazole. Next, the protein was eluted with 20 mM Tris (pH 8.0), 500 mM NaCl and 400 mM imidazole. The eluted fraction was concentrated and loaded onto a Superdex200 10/300 Increase size-exclusion column, equilibrated in buffer containing 20 mM Tris (pH 8.0) and 150 mM NaCl. Peak fractions were pooled, concentrated to 5 mg/ml using a centrifugal filter device (Millipore 10 kDa MW cutoff), and frozen in liquid nitrogen.

Nb35 was prepared as previously reported46,47. In brief, Nb35 was expressed in the periplasm of E. coli. The harvested cells were disrupted by sonication. Nb35 was purified by nickel affinity chromatography, followed by gel-filtration chromatography, and frozen in liquid nitrogen.

Constructs for expression of GPR30 and Gq

The human GPR30 gene (UniProtKB, Q99527) was subcloned into a modified pFastBac vector20, with an N-terminal haemagglutinin signal peptide followed by the Flag-tag epitope (DYKDDDDK) and the LgBiT fused to its C-terminus followed by a 3 C protease site and EGFP-His8 tag. A 15 amino sequence of GGSGGGGSGGSSSGG was inserted into both the N-terminal and C-terminal sides of LgBiT. Rat Gβ1 and bovine Gγ2 were subcloned into the pFastBac Dual vector. In detail, rat Gβ1 was cloned with a C-terminal HiBiT connected with a 15 amino sequence of GGSGGGGSGGSSSGG. Moreover, mini-Gsqi was subcloned into the C-terminus of the bovine Gγ2 with a nine amino sequence GSAGSAGSA linker. The resulting pFastBac dual vector can express the mini-Gsqi trimer.

Expression and purification of the human GPR30 – Gq complex

Earlier version: The recombinant baculovirus was prepared using the Bac-to-Bac baculovirus expression system (Thermo Fisher Scientific). For expression, 2 L of Sf9 cells at a density of 3 × 106 cells/mL were co-infected with baculovirus encoding GPR30 and mini-Gsqi trimer at the ratio of 1:1 and the cells were incubated at 30 ℃. After 48 hours, the collected cells were resuspended and dounce-homogenized in 20 mM Tris-HCl, pHC8.0 (4 ℃), 200 mM NaCl, 10% Glycerol, 200 mM NaHCO3, 25 mU/ml apyrase. The crude membrane fraction was collected by ultracentrifugation at 180,000g for 1Ch and solubilized in buffer, containing 20CmM Tris-HCl, pHC8.0, 150CmM NaCl, 1% n-dodecyl-beta-D-maltopyranoside (DDM) (Calbiochem), 0.2 % cholesteryl hemisuccinate (CHS) (Merck), 10% glycerol, 200 mM NaHCO3, and 25 mU/ml Apyrase, for 1 h at 4 ℃. The supernatant was separated from the insoluble material by ultracentrifugation at 180,000g for 20Cmin and incubated with 5 mL of M2 anti-flag affinity resin (Sigma) for 1 h at 4 ℃. The resin was washed with 20 column volumes of buffer containing 20CmM Tris-HCl, pHC8.0 500CmM NaCl, 10% Glycerol, 0.05% Lauryl Maltose Neopentyl Glycol (LMNG) (Antrace), 0.005% CHS and 200 mM NaHCO3. The complex was eluted in buffer containing 20CmM Tris-HCl, pHC8.0, 150CmM NaCl, 10% Glycerol, 0.01% LMNG, 0.001% CHS, 200 mM NaHCO3 and 0.2 mg/mL Flag peptide. The eluate was incubated with the Nb35 and scFv16 at 4 ℃. The complex was concentrated and purified by size exclusion chromatography on a Superose 6 increase (GE) column in 20CmM Tris-HCl, pH8.0, 150CmM NaCl, 0.01% LMNG, 0.001% CHS and 200 mM NaHCO3. Peak fractions were concentrated to 4.72 mg/ml for electron microscopy studies.

Revised version: The recombinant baculovirus was prepared using the Bac-to-Bac baculovirus expression system (Thermo Fisher Scientific). For expression, 2 L of Sf9 cells at a density of 3 × 106 cells/mL were co-infected with baculovirus encoding GPR30 and mini-Gsqi trimer at the ratio of 1:1 and the cells were incubated at 30°C. After 48 hours, the collected cells were resuspended and dounce-homogenized in 20 mM Tris-HCl, pH 8.0 (4°C), 200 mM NaCl, 10% Glycerol, 200 mM NaHCO3, 25 mU/ml apyrase. The crude membrane fraction was collected by ultracentrifugation at 180,000g for 1 h and solubilized in buffer, containing 20 mM Tris-HCl, pH 8.0, 150 mM NaCl, 1% n-dodecyl-beta-D-maltopyranoside (DDM) (Calbiochem), 0.2 % cholesteryl hemisuccinate (CHS) (Merck), 10% glycerol, 200 mM NaHCO3, and 25 mU/ml Apyrase, for 1 h at 4°C. The supernatant was separated from the insoluble material by ultracentrifugation at 180,000g for 20 min and incubated with 5 mL of M2 anti-flag affinity resin (Sigma) for 1 h at 4°C. The resin was washed with 20 column volumes of buffer containing 20 mM Tris-HCl, pH 8.0 500 mM NaCl, 10% Glycerol, 0.05% Lauryl Maltose Neopentyl Glycol (LMNG) (Antrace), 0.005% CHS and 200 mM NaHCO3. The complex was eluted in buffer containing 20 mM Tris-HCl, pH 8.0, 150 mM NaCl, 10% Glycerol, 0.01% LMNG, 0.001% CHS, 200 mM NaHCO3 and 0.2 mg/mL Flag peptide. The eluate was incubated with the Nb35 and scFv16 at 4°C. The complex was concentrated and purified by size exclusion chromatography on a Superose 6 increase (GE) column in 20 mM Tris-HCl, pH8.0, 150 mM NaCl, 0.01% LMNG, 0.001% CHS and 200 mM NaHCO3. Peak fractions were concentrated to 4.72 mg/ml for electron microscopy studies.

Sample vitrification and cryo-EM data acquisition

The purified complex was applied onto a freshly glow-discharged Quantifoil Au grid (R1.2/1.3, 300 mesh), and plunge-frozen in liquid ethane by using a Vitrobot Mark IV. Data collections were performed on a 300kV Titan Krios G3i microscope (Thermo Fisher Scientific) equipped with a BioQuantum K3 imaging filter and a K3 direct electron detector (Gatan).

Earlier version: A total of 9824 movies were acquired with a calibrated pixel size of 0.83CÅCpix-1 and with a defocus range of -0.8 to -1.6Cμm, using EPU. Each movie was acquired for 2.6Cs and split into 48 frames, resulting in an accumulated exposure of about 50.660Celectrons per Å2.

Revised version: A total of 9824 movies were acquired with a calibrated pixel size of 0.83 Å pix-1 and with a defocus range of -0.8 to -1.6 μm, using EPU. Each movie was acquired for 2.6 s and split into 48 frames, resulting in an accumulated exposure of about 50.660 electrons per Å2.

Image processing

All acquired dose-fractionated movies were imported into CryoSPARC v4.448 and subjected to beam-induced motion correction. The contrast transfer function (CTF) parameters were estimated and a total of 10,148,422 particles were extracted. The particles were subjected to 2D classifications, Ab-initio reconstruction and several rounds of hetero refinement and Non-uniform refinement. Next, the particles were subjected to 3D classification with a mask on the receptor. Then the particle sets were subjected to Reference Based Motion Correction. Motion-corrected 522,404 particles were subjected to Non-uniform refinement, yielding a map with a global nominal resolution of 3.15 Å, with the gold standard Fourier Shell Correlation (FSC=0.143) criteria49. Moreover, the 3D model was refined with a mask on the receptor. As a result, the receptor has a higher resolution with a nominal resolution of 3.18 Å. The overall and receptor focused maps were combined by phenix50. The processing strategy is described in Supplementary Figure 2.

Model building and refinement

The density map was sharpened by phenix.auto_sharpen51 and the quality of the density map was sufficient to build a model manually in COOT52,53. The model building was facilitated by the Alphafold-predicted structure. We manually fitted GPR30, the Gq heterotrimer and scFv16 into the map. We then manually readjusted the model into using COOT and refined it using phenix.real_space_refine50,54 (v.1.19) with the secondary-structure restraints using phenix secondary_structure_restraints.

Vector construction and transfection

Human Gpr30 cDNA was obtained from human hepatoblastoma-derived HepG2 cells. The coding sequences of human Gpr30 were inserted into the multi-cloning site of the plasmid vector pCXN2, which was generated in our laboratory via modification of pCAGGS, between the KpnI and EcoRI sites. The N- and C-terminal HA-tagged Gpr30 was amplified using a reverse primer containing the HA sequenceOne amino acid mutation of human GPR30 was generated as follows: the targeted amino acid was changed to alanine (GCC) using a two-step PCR method with the QuikChange® Primer Design Program by Agilent (https://www.agilent.com/store/primerDesignProgram.jsp), and the coding sequence with each mutation was inserted into the multi-cloning site of pCXN2 between the KpnI and EcoRI sites. The mutations generated were D111A, E121A, R122A, D125A, S134A, Q138A, D210A, Q215A, E218A, N276A, Q296A, C207A, P71A, and H307A.

Cell line sources and transfection

Earlier version: HEK293A cells (female origin; Thermo Fisher Scientific) were maintained at 37 ℃ and 5% CO2. These vectors were transfected using the lipofection method (Lipofectamine™ 2000 Transfection Reagent, 11668019, Invitrogen). Stable cell lines were established through antibiotic drug selection. The HEK293A cells were transfected with the N-terminal HA-tagged wild-type and mutant Gpr30 and maintained in the culture medium containing 1 mg/ml of G418. After 2-3 weeks of antibiotic drug selection, stable expression of exogenous GPR30 was confirmed by three independent flow cytometric analyses.

Revised version: HEK293A cells (female origin; Thermo Fisher Scientific) were maintained at 37°C and 5% CO2. These vectors were transfected using the lipofection method (Lipofectamine 2000 Transfection Reagent, 11668019, Invitrogen). Stable cell lines were established through antibiotic drug selection. The HEK293A cells were transfected with the N-terminal HA-tagged wild-type and mutant Gpr30 and maintained in the culture medium containing 1 mg/ml of G418. After 2-3 weeks of antibiotic drug selection, stable expression of exogenous GPR30 was confirmed by three independent flow cytometric analyses.

TGFα shedding assay

Earlier version: The TGFα shedding assay was performed according to a previously published protocol26. The AP-TGFα expression vector was provided by Dr. Inoue and Dr. Aoki, Tohoku University. HEK293 cells were seeded in 12-well plates at a density of 1×105 cells/well and cultured for 24 h. At 70% confluency, a mixture of plasmid vectors containing GPCR (see ‘Vector construction and transfection’ section) and AP-TGFα was transfected into the cells using Lipofectamine 2000 transfection reagent. After another 24 h of incubation, the cells were detached with 0.05% trypsin/EDTA (32777-44, Nacalai Tesque), suspended in Hanks’ balanced salt solution, and seeded in 96-well plates. The cells were stimulated for 1 h with 1–44×10-3 M (final concentration) of NaHCO3 at 37 ℃, under 0.03% CO2. Conditioned media (CM) was transferred to another plate, and 80 µl of alkaline phosphatase (AP) solution (40 mM Tris-HCl, pH 9.5, 40 mM NaCl, 10 mM MgCl2, 10 mM p-nitrophenylphosphate disodium salt hexahydrate) was added to both plates, which were then incubated at 37 ℃. The optical density at 405 nm (OD405) was measured using a microplate reader (iMark, Bio-Rad) at 5 min, 30 min, 1 h, and 2 h, depending on the reaction rate. The percentage of shed AP-TGF was calculated using the following equations:

Revised version: The TGFα shedding assay was performed according to a previously published protocol26. The AP-TGFα expression vector was provided by Dr. Inoue and Dr. Aoki, Tohoku University. HEK293 cells were seeded in 12-well plates at a density of 1×105 cells/well and cultured for 24 h. At 70% confluency, a mixture of plasmid vectors containing GPCR (see ‘Vector construction and transfection’ section) and AP-TGFα was transfected into the cells using Lipofectamine 2000 transfection reagent. After another 24 h of incubation, the cells were detached with 0.05% trypsin/EDTA (32777-44, Nacalai Tesque), suspended in Hanks’ balanced salt solution, and seeded in 96-well plates. The cells were stimulated for 1 h with 1–44×10-3 M (final concentration) of NaHCO3 at 37°C, under 0.03% CO2. Conditioned media (CM) was transferred to another plate, and 80 µl of alkaline phosphatase (AP) solution (40 mM Tris-HCl, pH 9.5, 40 mM NaCl, 10 mM MgCl2, 10 mM p-nitrophenylphosphate disodium salt hexahydrate) was added to both plates, which were then incubated at 37°C. The optical density at 405 nm (OD405) was measured using a microplate reader (iMark, Bio-Rad) at 5 min, 30 min, 1 h, and 2 h, depending on the reaction rate. The percentage of shed AP-TGF was calculated using the following equations:

Added in the revised version: AP activity = ΔOD405 (1 – 0 h)

Added in the revised version: % CM (conditioned media) = APCM / (APCM + APCell)

Calcium assay

Earlier version: Stable HEK293 cells (2×104 cells/well) were seeded into a black wall and clear bottom 96-well plate 24 h before the assay. Cells at 90–100% confluency were incubated with HEPES buffer (1× HBSS, 2.5 mM probenecid, 25 mM HEPES, pH 7.4) containing 10 µM Fluo-8 AM (21080, AAT-Bio) for 60 min at 37 ℃ and 5% CO2. The cells were washed twice with HEPES buffer. The cells were stimulated with indicated concentrations of NaHCO3 and ATP at 37 ℃ and 0.03% CO2. The fluorescence intensity was analysed using FlexStation 3 (Molecular Devices, Ex/Em = 490/525 nm).

Revised version: Stable HEK293 cells (2×104 cells/well) were seeded into a black wall and clear bottom 96-well plate 24 h before the assay. Cells at 90–100% confluency were incubated with HEPES buffer (1× HBSS, 2.5 mM probenecid, 25 mM HEPES, pH 7.4) containing 10 µM Fluo-8 AM (21080, AAT-Bio) for 60 min at 37°C and 5% CO2. The cells were washed twice with HEPES buffer. The cells were stimulated with indicated concentrations of NaHCO3 and 50 µM ATP at 37°C and 0.03% CO2. The fluorescence intensity was analysed using FlexStation 3 (Molecular Devices, Ex/Em = 490/525 nm). The difference between the maximum and minimum fluorescent values during 15 to 60 sec was normalized by those with ATP stimulation. Nonlinear regression (four-parameter) was used for curve-fitting, and EC50 and Top (maximum response) were calculated.

Flow cytometry

Earlier version: HEK293A cells stably expressing the N-terminal HA-tagged wild-type and mutant GPR30 were incubated with the calcium-free Minimum Essential Medium (S-MEM, #11380037, Gibco™) for 30 min at 37 ℃ and 5.0% CO2 before the assay. The cells were blocked with PBS containing 2% goat serum for 10 min at 4 ℃, incubated with the biotin-conjugated anti-HA antibody (2.5 µg/ml, #12158167001, Roche) for 60 min at 4 ℃, and then stained with Alexa Fluor 488-conjugated anti-rat IgG (5 µg/ml, #A-11006, Thermo Fisher Scientific) for 40 min at 4 ℃. LSRFortessa (BD Biosciences) was used for analysis. For intracellular staining, the cells were fixed and stained using eBioscience™ Fixable Viability Dye eFluor 780 (#65-0865-14, Thermo Fisher Scientific) and eBioscience™ Fox3/Transcription Factor Staining Buffer Set (#00-5523-00, Thermo Fisher Scientific).

Revised version: HEK293A cells stably expressing the N-terminal HA-tagged wild-type and mutant GPR30 were incubated with the calcium-free Minimum Essential Medium (S-MEM, #11380037, Gibco) for 30 min at 37°C and 5.0% CO2 before the assay. The cells were blocked with PBS containing 2% goat serum for 10 min at 4°C, incubated with the biotin-conjugated anti-HA antibody (2.5 µg/ml, #12158167001, Roche) for 60 min at 4°C, and then stained with Alexa Fluor 488-conjugated anti-rat IgG (5 µg/ml, #A-11006, Thermo Fisher Scientific) for 40 min at 4°C. LSRFortessa (BD Biosciences) was used for analysis. For intracellular staining, the cells were fixed and stained using eBioscience Fixable Viability Dye eFluor 780 (#65-0865-14, Thermo Fisher Scientific) and eBioscience Fox3/Transcription Factor Staining Buffer Set (#00-5523-00, Thermo Fisher Scientific).

Analysis of cell surface expression of GPR30

Earlier version: The cell surface expression of wild-type and mutant hGPR30-HA was analysed via cell surface protein isolation using a Cell Surface Protein Isolation Kit (#89881, Thermo Scientific™) followed by western blotting. Briefly, HEK293 cells transiently expressing wild-type or mutant hGPR30-HA were biotinylated for 30 min at 4 ℃. Then, the cells were harvested, and biotinylated proteins were isolated with avidin binding.

Revised version: The cell surface expression of wild-type and mutant hGPR30-HA was analysed via cell surface protein isolation using a Cell Surface Protein Isolation Kit (#89881, Thermo Scientific) followed by western blotting. Briefly, HEK293 cells transiently expressing wild-type or mutant hGPR30-HA were biotinylated for 30 min at 4°C. Then, the cells were harvested, and biotinylated proteins were isolated with avidin binding.

Western blotting

Earlier version: Cell lysates were separated using sodium dodecyl sulfate-polyacrylamide gel electrophoresis under reducing conditions and transferred to polyvinylidene difluoride membranes (Immobilon P, IPVH00010, Millipore). Primary antibodies used were anti-HA High Affinity, 1:1000 dilution, 11867423001, Roche; Na-K-ATPase, 1:1000 dilution, #3010, Cell Signaling Technology; β-Actin (AC-15), 1:1000 dilution, sc-69879, Santa Cruz Biotechnology. Secondary antibodies used were anti-rabbit IgG, HRP-linked Antibody, 1:5000 dilution, #7074, Cell Signaling Technology; anti-mouse IgG, HRP-linked Antibody, 1:5000 dilution, #7076, Cell Signaling Technology; anti-rat IgG, HRP-Linked Whole Ab Goat, 1:5000 dilution, NA935, Cytiva. The membranes were probed at 4 ℃ overnight with the primary antibodies. The membranes were subsequently incubated with the corresponding secondary antibodies. The signals were detected with ECL Prime (RPN2236, Cytiva) or ImmunoStar LD (296-69901, FUJIFILM Wako Pure Chemical Corporation) using a chemiluminescence imaging system (CFusion FX7, Vilber).

Revised version: Cell lysates were separated using sodium dodecyl sulfate-polyacrylamide gel electrophoresis under reducing conditions and transferred to polyvinylidene difluoride membranes (Immobilon P, IPVH00010, Millipore). Primary antibodies used were anti-HA High Affinity, 1:1000 dilution, 11867423001, Roche; Na-K-ATPase, 1:1000 dilution, #3010, Cell Signaling Technology; β-Actin (AC-15), 1:1000 dilution, sc-69879, Santa Cruz Biotechnology. Secondary antibodies used were anti-rabbit IgG, HRP-linked Antibody, 1:5000 dilution, #7074, Cell Signaling Technology; anti-mouse IgG, HRP-linked Antibody, 1:5000 dilution, #7076, Cell Signaling Technology; anti-rat IgG, HRP-Linked Whole Ab Goat, 1:5000 dilution, NA935, Cytiva. The membranes were probed at 4°C overnight with the primary antibodies. The membranes were subsequently incubated with the corresponding secondary antibodies. The signals were detected with ECL Prime (RPN2236, Cytiva) or ImmunoStar LD (296-69901, FUJIFILM Wako Pure Chemical Corporation) using a chemiluminescence imaging system (CFusion FX7, Vilber).